(assume a0 (or (< (+ x y) 1) (< 3 x)))
(assume a2 (= 0 y))
(assume a1 (= x 2))
(step t1 (cl (< (+ x y) 1) (< 3 x)) :rule or :premises (a0))
(step t2 (cl (not (< 3 x)) (not (= x 2))) :rule la_generic :args (1/1 -1/1))
(step t3 (cl (not (< 3 x))) :rule resolution :premises (a1 t2))
(step t4 (cl (< (+ x y) 1)) :rule resolution :premises (t1 t3))
(step t5 (cl (not (< (+ x y) 1)) (not (= x 2)) (not (= 0 y))) :rule la_generic :args (1/1 1/1 -1/1))
(step t6 (cl) :rule resolution :premises (t5 t4 a1 a2))